---
  title: "2"
  description: ""

---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 68be035dca   
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```



`@sct`

```{python}

```







---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 20146471f5   
```

dasdasdas


`@hint`






`@possible_answers`


`@feedbacks`




