---
  title: "5"
  description: ""

---
## Insert exercise title here

```yaml
type: VideoExercise 
xp: 50 
key: 2a27a3ad98   
```

`@projector_key`
767f978124eacfb58299b201128d0cab

---
## Insert exercise title here

```yaml
type: NormalExercise 
xp: 100 
key: f2c2776113   
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```

`@sample_code`

```{python}

```

`@solution`

```{python}

```

`@sct`

```{python}

```







---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: 684e26dbdc   
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```



`@sct`

```{python}

```







---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 2d3f0a0cc3   
```




`@hint`






`@possible_answers`


`@feedbacks`






---
## Insert exercise title here

```yaml
type: TabExercise 
xp: 100 
key: 281e755569   
```





`@pre_exercise_code`

```{python}

```

`@sample_code`

```{python}

```








***



```yaml
type: NormalExercise 
xp: 50 
key: 5e6397b3dd   
```



`@instructions`


`@hint`



`@sample_code`

```{python}

```

`@solution`

```{python}

```

`@sct`

```{python}

```








***



```yaml
type: MultipleChoiceExercise 
xp: 50 
key: b15c47549d   
```



`@instructions`


`@hint`





`@sct`

```{python}

```








---
## Insert exercise title here

```yaml
type: BulletExercise 
xp: 100 
key: f9b939de73   
```





`@pre_exercise_code`

```{python}

```

`@sample_code`

```{python}

```








***



```yaml
type: NormalExercise 
xp: 50 
key: 8310a6dc12   
```



`@instructions`


`@hint`



`@sample_code`

```{python}

```

`@solution`

```{python}

```

`@sct`

```{python}

```








***



```yaml
type: NormalExercise 
xp: 50 
key: e92f87f179   
```



`@instructions`


`@hint`



`@sample_code`

```{python}

```

`@solution`

```{python}

```

`@sct`

```{python}

```






