---
  title: "3"
  description: ""

---
## Insert exercise title here

```yaml
type: PureMultipleChoiceExercise 
xp: 50 
key: 24b4d0d0b8   
```




`@hint`






`@possible_answers`


`@feedbacks`






---
## Insert exercise title here

```yaml
type: MultipleChoiceExercise 
xp: 50 
key: a2ec612666   
```



`@instructions`


`@hint`


`@pre_exercise_code`

```{python}

```



`@sct`

```{python}

```





